<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 125184  -->
  <question type="stack">
    <name>
      <text>System of Equations with No Solution</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Give two equations in the form \(ax+by=c\) such that when you try to solve them simultaneously there are no solutions.
</p>
<p> Equation 1: [[input:ans1]] [[validation:ans1]]</p>
<p> Equation 2: [[input:ans2]] [[validation:ans2]]</p>
[[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b><h3>Worked solution:</h3></b> </p>
<p>Here is an example of  two equations such that when you try to solve them simultaneously there are no solutions:</p>
<p> \({@eq1@} \quad \) Equation (1) </p>
<p> \({@eq2@}\quad \) Equation (2) </p>
<p>Indeed, when you try to solve simultaneously by subtracting equation (1) from equation (2), you get \({@eq2-eq1@}. \) </p>
<p></p> This is not true, so there are no solutions.<p> 
</p>Graphically, if we plot these two straight lines, we can see that there are no solutions because the lines are parallel.<p></p>
<p>{@plot([line1,line2],[x,-4,4],[y,-2,5], [box, false], [yx_ratio, 1], [axes, solid], [xtics, -3, 1, 3],[ytics, -1, 1, 4],grid2d,[legend,"Equation 1", "Equation 2"])@}</p>
<p>Any two lines that are parallel will have no solutions.</p>
<p>Two lines are parallel if they have the same gradient. If the coefficients of \(x\) and \(y\) are the same in both equations, as in (1) and (2) above, then the gradients will be the same. The same holds true also when the coefficients of \(x\) and \(y\) are  are in the same ratio, for example</p>
<p> \({@eq1@}\) </p>
<p>\({@expand(3*lhs(eq2))=rhs(eq1)+1@}\).</p>
<p>Note also that the constants in the two equations should not be equal or in the same ratio as the coefficients of \(x\) and \(y\). Indeed in this case the two equations would correspond to the same line and there would be infinitely many solutions.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*define coefficients of the two lines for the worked solution*/
aa:rand(3)+1;
bb:rand(3)+1;
cc_1:(rand(3)+1);
cc_2:rand_with_prohib(1,3,[cc_1]); /*should not be the same line*/

eq1: aa*x+bb*y=bb*cc_1;  /*want integer intercept*/
eq2: aa*x+bb*y=bb*cc_2;
line1: expand(-aa/bb*x+cc_1);
line2: expand(-aa/bb*x+cc_2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Displayed solutions:</p>
<p>{@eq1@}</p>
<p>{@eq2@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">NA</p>
<p id="reviewer">Godfrey  Wabwire</p>
<p id="reviewer">Juma Zevick</p>
<p id="description">This question asks for two equations in the form ax + by = c that cannot be solved simultaneously, resulting in no solution.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>eq1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>eq2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[corr_form_1:if (lhs(ans1) - coeff(lhs(ans1),x)*x - coeff(lhs(ans1),y)*y =0 and ({hipow(lhs(ans1),x),hipow(lhs(ans1),y)}={1} and {hipow(rhs(ans1),x),hipow(rhs(ans1),y)}={0}))  then 1 else 0;
corr_form_2:if (lhs(ans2) - coeff(lhs(ans2),x)*x - coeff(lhs(ans2),y)*y =0 and ({hipow(lhs(ans2),x),hipow(lhs(ans2),y)}={1} and {hipow(rhs(ans2),x),hipow(rhs(ans2),y)}={0}))  then 1 else 0;


plot_form_ans1:expand(solve(ans1,y)[1])
plot_form_ans2:expand(solve(ans2,y)[1])

plot_ans1:rhs(plot_form_ans1);
plot_ans2:rhs(plot_form_ans2);

grad_1:coeff(plot_ans1,x);
grad_2:coeff(plot_ans2,x);

xmin:if solve([ans1,ans2],[x,y])=[] then -4 elseif rhs(solve([ans1,ans2],[x,y])[1][1])<=-4 then rhs(solve([ans1,ans2],[x,y])[1][1])-1 else -4;
xmax:if solve([ans1,ans2],[x,y])=[] then 4 elseif rhs(solve([ans1,ans2],[x,y])[1][1])>=4 then rhs(solve([ans1,ans2],[x,y])[1][1])+1 else 4;

ymin:if solve([ans1,ans2],[x,y])=[] then -2 elseif rhs(solve([ans1,ans2],[x,y])[1][2])<=-2 then rhs(solve([ans1,ans2],[x,y])[1][2])-1 else -2;
ymax:if solve([ans1,ans2],[x,y])=[] then 5 elseif rhs(solve([ans1,ans2],[x,y])[1][2])>=5 then rhs(solve([ans1,ans2],[x,y])[1][2])+1 else 5;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>grad_1</sans>
        <tans>grad_2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Good thinking, but there are actually infinitely many solutions to these simultaneous equations. Graphically they correspond to the equation of the same line.</p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for incorrect answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for incorrect answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>corr_form_1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.25</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Equation 1 is not in the correct form.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Checks for incorrect answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>corr_form_2</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>0</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.25</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>0</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Equation 2 is not in the correct form.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1876887423</deployedseed>
    <deployedseed>355244185</deployedseed>
    <deployedseed>1568932832</deployedseed>
    <deployedseed>820217035</deployedseed>
    <deployedseed>1617883325</deployedseed>
    <deployedseed>1971640918</deployedseed>
    <deployedseed>2145720668</deployedseed>
    <deployedseed>2022278558</deployedseed>
    <deployedseed>1498250775</deployedseed>
    <deployedseed>1612572241</deployedseed>
    <deployedseed>79443499</deployedseed>
    <deployedseed>706324495</deployedseed>
    <deployedseed>143900811</deployedseed>
    <deployedseed>910660833</deployedseed>
    <deployedseed>442686895</deployedseed>
    <deployedseed>1161136102</deployedseed>
    <deployedseed>1510577845</deployedseed>
    <deployedseed>1326808790</deployedseed>
    <deployedseed>479054911</deployedseed>
    <deployedseed>2139026926</deployedseed>
    <deployedseed>1599138748</deployedseed>
    <deployedseed>41264433</deployedseed>
    <deployedseed>950072879</deployedseed>
    <deployedseed>629224167</deployedseed>
    <deployedseed>142252389</deployedseed>
    <deployedseed>346199171</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>eq1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>eq2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
    <tags>
      <tag><text>Bridging</text>
</tag>
      <tag><text>Bridging Week 1</text>
</tag>
      <tag><text>Open Bridging</text>
</tag>
    </tags>
  </question>

</quiz>