<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 140566  -->
  <question type="stack">
    <name>
      <text>Planes curves and Parametric Equations: Question 4</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the equation of the tangent line \((T)\) to the curve \({@eq_curv@}\) at the point \(P({@pt_x@},{@pt_y@})\).</p>
<p>\((T):\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3><b>Worked solution</b></h3>
<p>Find the equation of the tangent line \((T)\) to the curve \({@eq_curv@}\) at the point \(P({@pt_x@},{@pt_y@})\).</p>
<p>Since \( \displaystyle \frac{dy}{dx} ={@dif_y@} \) . Then  \( \displaystyle  \frac{dy}{dx} \vert_{({@x=pt_x@})}={@val_dif@} \)  <p>
<p> Using: \( \displaystyle \frac{\Delta y}{ \Delta x} = \displaystyle {@val_dif@} \) we have as follows:
<p>\(
\begin{align}
\displaystyle \frac{\Delta y}{ \Delta x} = {@val_dif@} & \implies & \displaystyle \frac{y - y_{1}}{ x - x_{1}} = {@val_dif@}\\
& \implies & {@y-pt_y=val_dif*(x-pt_x)@}, \text{because } x_1={@pt_x@},\,  y_1={@pt_y@};    \\
& \implies & {@y-pt_y@}={@expand(val_dif*(x-pt_x))@} \\
& \implies & {@ta_1@} \\
\end{align}
\)</p>
 <p>Thus \((T): \displaystyle {@ta_1@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Define the randomisation coefficients*/
aa: rand_with_prohib(-5,5,[0]);
bb: rand_with_prohib(-5,5,[0]);

/*Equation of the curve*/
eq_curv: aa*y=bb*x^2; 

/*Coordinates of the point P*/
pt_x: rand_with_step(-5,5,1);
pt_y:  rand_with_step(-5,5,1);

/*Find Equation of the tangent line*/
expr_y: (bb/aa)*x^2; /*Redefine y=f(x)*/
dif_y: diff(expr_y,x); /* Find dy/dx*/
val_dif: at(dif_y, [x=pt_x]);

/*Model answer*/
ta_1: expand(y=val_dif*(x-pt_x)+pt_y);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Curve: \({@eq_curv@}\) and \((T): {@ta_1@}\)</p> ]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">n/a</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer"></p>
<p id="description">The question focuses on finding the equation finding the equation of the tangent line \((T)\) to the curve \({@eq_curv@}\) at the point \(P({@pt_x@},{@pt_y@})\) </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>12</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You should use the formula \(y - y_{1}= a \cdot (x - x_{1})\), where \(a\) is called the slope and stand for \( \frac{\Delta y}{\Delta x}\) at \(x={@pt_x@}\)</p>
<p> You should find \((T): {@ta_1@}\). Check the worked solution.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1494393360</deployedseed>
    <deployedseed>170376958</deployedseed>
    <deployedseed>199781999</deployedseed>
    <deployedseed>2079879708</deployedseed>
    <deployedseed>646365751</deployedseed>
    <deployedseed>1724353357</deployedseed>
    <deployedseed>416566863</deployedseed>
    <deployedseed>622968742</deployedseed>
    <deployedseed>1465429462</deployedseed>
    <deployedseed>2051070642</deployedseed>
    <deployedseed>229167435</deployedseed>
    <deployedseed>1612560259</deployedseed>
    <deployedseed>770462388</deployedseed>
    <deployedseed>1158125434</deployedseed>
    <deployedseed>2030060713</deployedseed>
    <deployedseed>915079147</deployedseed>
    <deployedseed>1248551269</deployedseed>
    <deployedseed>1351470551</deployedseed>
    <deployedseed>893080576</deployedseed>
    <deployedseed>1448698283</deployedseed>
    <deployedseed>413643815</deployedseed>
    <deployedseed>2062302007</deployedseed>
    <deployedseed>1208919340</deployedseed>
    <deployedseed>136309633</deployedseed>
    <deployedseed>7988865</deployedseed>
    <deployedseed>671697973</deployedseed>
    <deployedseed>1027732334</deployedseed>
    <deployedseed>791213161</deployedseed>
    <deployedseed>1396707268</deployedseed>
    <deployedseed>394715216</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>