<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 143670  -->
  <question type="stack">
    <name>
      <text>Planes curves and Parametric Equations: Question 1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p> Let </p>
<p>\((E)=\{(x, y): {@xx@} \, \mbox{and} \, {@yy@}, t \in \mathbb{R} \} \).</p>
<p> This means: \((E)\) represent the set of points \((x, y)\) such that \({@xx@}\) and \({@yy@}, t \in \mathbb{R}\).</p>
<br>1.  Determine the equation of \((E)\).
<br> \( (E):\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
<br>2. Determine the nature of \((E)\) in the following cases:
<br> a) For \(t={@tt@}\), Nature:  [[input:ans2]] [[validation:ans2]] [[feedback:prt2]].
<br> b) For \(t \in \mathbb{R}\), Nature:  [[input:ans3]] [[validation:ans3]] [[feedback:prt3]].
<br> 3. Sketch the plane curve \((E)\): take \(0 \leq t \leq {@rand([2,4,6,8,10])@} \cdot \pi\). </p>
<p><b>Hint:</b> For sketching the plane, you use the sliders 'a' and 'b'.</p>
<center>
[[jsxgraph width="500px" height="500px" input-ref-ansa='ansaRef'
 input-ref-ansb='ansbRef']]               
            const board = JXG.JSXGraph.initBoard(divid, {
            boundingbox: [-6, 6, 6, -6], axis: true
        });

        var b = board.create('slider', [[-5,4.5],[-2,4.5],[-2,{#a_init#},4]],{name:'b', snapWidth: 1});      
        var a = board.create('slider', [[-5,5],[-2,5],[0,{#a_init#},5]],{name:'a', snapWidth: 1});
        var c = board.create('curve', [function(phi){return a.Value()+b.Value()*phi; }, [0, 0],0, 8*Math.PI],
             {curveType:'polar', strokewidth:4});     

        stack_jxg.bind_slider(ansaRef,a); 
        stack_jxg.bind_slider(ansbRef,b); 

       
        board.update();
 
[[/jsxgraph]]
</center>
<p style="visibility:hidden">[[input:ansa]] [[validation:ansa]] [[feedback:prta]]</p>
<p style="visibility:hidden">[[input:ansb]] [[validation:ansb]] [[feedback:prtb]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Let \((E)=\{(x, y): {@xx@} \, \mbox{and} \, {@yy@}, t \in \mathbb{R} \} \), which means that</p>
<p> \((E)\) represent the set of points \((x, y)\) such that \({@xx@}\) and \({@yy@}, t \in \mathbb{R}\).</p>
<p>1.  Determine the equation \(E\)</p>
<p>We need to eliminate \(t\) as follows:</p>
<p>\({@xx@} \implies {@xx^2@} \) </p>
<p>\({@yy@} \implies {@yy^2@} \) </p>
<p>Then  \(x^2+ y^2=\left({@bb^2*cos(aa*t)^2@} \right)+\left({@bb^2*sin(aa*t)^2@}\right)=1\). </p>
<p> Recall that \(\forall x \in \mathbb{R}, {@cos(x)^2+ sin(x)^2=1@}\)</p>
<p>\( x^2+y^2={@bb^2@} \cdot ({@cos(aa*t)^2@}+{@sin(aa*t)^2@})={@bb^2@} \)</p>
<p> Therefore \( (E): x^2+ y^2={@bb^2@} \)</p>
<p>2. Determine the nature of \((E)\) in the following cases:</p>
<p> a) For \(t={@tt@}\), we need to substitute the value of \(t\) in the expressions of \(x\) and \(y\). </p>
<p> So, we have a point with coordinates \((x({@tt@}), y({@tt@}))={@pt@}\).</p>
<p> b) For \(t \in \mathbb{R}\), we have \( (E):x^2+ y^2={@bb^2@}\), then a circle of radius {@bb@}.</p>
<p> 3. Sketch the plane curve: take \(0 \leq t \leq {@rand([2,4,6,8,10])@} \cdot \pi\). </p>
<p><b>Hint:</b> For sketching, you use the sliders 'a' and 'b'.</p>
<centre>
[[jsxgraph]]
               
            const board = JXG.JSXGraph.initBoard(divid, {
            boundingbox: [-6, 6, 6, -6], axis: true
        });
             
        var a = board.create('slider', [[-5,5],[-2,5],[0,{#a_cor#},5]],{name:'a'});
        var b = board.create('slider', [[-5,4.5],[-2,4.5],[0,{#b_cor#},4]],{name:'b'});      
        var c = board.create('curve', [function(phi){return a.Value()+b.Value()*phi; }, [0, 0],0, 8*Math.PI],
             {curveType:'polar', strokewidth:4});     
 
[[/jsxgraph]]
</centre>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*Randomising the coefficients of the parametric curves*/
aa: rand(5)+1;
bb: rand_with_step(1,4,1);
/*Define the parametric curves*/
xx: x=bb*cos(aa*t);
yy: y=bb*sin(aa*t);
/*Randomised the time t*/
tt: rand([0, %pi, %pi/2, %pi/4, %pi/6, -%pi/2]);
/*Coordinates of the point at t*/
pt_x: at(bb*cos(aa*t), [t=tt]);
pt_y: at(bb*sin(aa*t), [t=tt]);
pt: oo(pt_x, pt_y);



/*Model answers*/
ta_1: x^2+y^2=bb^2;
ta_2: random_permutation([["Circle", false],["Point", true],["Line", false], ["tangent", false], ["Parabola", false]]);
ta_3: random_permutation([["Circle", true],["Point", false],["Line", false], ["tangent", false], ["Parabola", false]]);


/*jsxgraph*/
/*Values of parameters by default*/
a_init: 1/5; 
b_init: 1/5;

/*Correct values of a and b*/
a_cor: bb; 
b_cor: 0;


]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[ <p>\( (E):{@ta_1@}\)</p>
<p> a) For \(t={@tt@}\), Nature: {@mcq_correct(ta_2)[1]@}.</p>
<p> b) For \(t \in \mathbb{R}\), Nature:{@mcq_correct(ta_3)[1]@}.</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">This question requires the student to determine the equation of the parametrically defined plane curve \((E)\) given by \((x, y) = (@xx@, @yy@)\) for \(t \in \mathbb{R}\). Additionally, the student needs to explore the nature of \((E)\) for specific values of \(t\) and sketch the curve using JSXGraph with provided sliders 'a' and 'b'.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansa</name>
      <type>algebraic</type>
      <tans>a_cor</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansb</name>
      <type>algebraic</type>
      <tans>b_cor</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p> To find the equation, you need to eliminate the parameter \(t\) by square each term and do their summation.</p>
<p> Recall that \(\forall x \in \mathbb{R}, {@cos(x)^2+ sin(x)^2=1@}\)</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_2)[1]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p> Know that for each single value of \(t_{0}\), we will have at point having a coordinate \(  (x(t_{0}), y(t_{0})) \) </p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta_3)[1]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p> For \(t \in \mathbb{R}\), the point \((x,y)\) will move by describing a circle of radius 1 and centre the origin </p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prta</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansa</sans>
        <tans>a_cor</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prta-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[The value of <b>a</b> you chose is <b>{@ansa@}</b>, which is correct.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prta-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[The value of <b>a</b> you chose is <b>{@ansa@}</b>, which is wrong. The correct is {@a_cor@}.]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtb</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansb</sans>
        <tans>b_cor</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtb-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[The value of <b>b</b> you chose is <b>{@ansb@}</b>, which is correct.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtb-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[The value of <b>b</b> you chose is <b>{@ansb@}</b>, which is wrong. The correct is {@b_cor@}.]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>923673307</deployedseed>
    <deployedseed>920633837</deployedseed>
    <deployedseed>531212504</deployedseed>
    <deployedseed>706038680</deployedseed>
    <deployedseed>1360725588</deployedseed>
    <deployedseed>706581349</deployedseed>
    <deployedseed>988369858</deployedseed>
    <deployedseed>668729233</deployedseed>
    <deployedseed>173946901</deployedseed>
    <deployedseed>854304650</deployedseed>
    <deployedseed>729158799</deployedseed>
    <deployedseed>1899506209</deployedseed>
    <deployedseed>1326210311</deployedseed>
    <deployedseed>739182921</deployedseed>
    <deployedseed>511833188</deployedseed>
    <deployedseed>2128734667</deployedseed>
    <deployedseed>905856107</deployedseed>
    <deployedseed>492377647</deployedseed>
    <deployedseed>263594258</deployedseed>
    <deployedseed>1504181159</deployedseed>
    <deployedseed>7181174</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta_3))</value>
      </testinput>
      <testinput>
        <name>ansa</name>
        <value>a_cor</value>
      </testinput>
      <testinput>
        <name>ansb</name>
        <value>b_cor</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prta</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prta-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtb</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtb-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>